
## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.
Instructions

- **Download** this notebook as you would any other ipynb file 
- **Upload** to Google Colab or work locally (if you have that set-up)
- **Delete** `raise NotImplementedError()`

- **Write** your code in the `# YOUR CODE HERE` space


- **Execute** the Test cells that contain assert statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)

- **Save** your notebook when you are finished
- **Download** as a ipynb file (if working in Colab)
- **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)



# Sprint Challenge
## *Data Science Unit 4 Sprint 1*

After a week of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). I've sampled the data for you to something more managable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge. As you work on the challenge, I suggest adding notes about your findings and things you want to analyze in the future.

## Challenge Objectives
Successfully complete all these objectives to earn full credit. 

**Successful completion is defined as passing all the unit tests in each objective.**  

Each unit test that you pass is 1 point. 

There are 5 total possible points in this sprint challenge. 


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernal

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

3) Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section). 
____



### Import Data

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
df = pd.read_json(data_url, lines=True)
df.head()
```

Output:
```
business_id  cool                date  funny  \
0  nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1  eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2  6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3  k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4  6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   

                review_id  stars  \
0  eZs2tpEJtXPwawvHnHZIgQ      1   
1  DoQDWJsNbU0KL1O29l_Xug      4   
2  DDOdGU7zh56yQHmUnL1idQ      3   
3  LfTMUWnfGFMOfOIyJcwLVA      1   
4  zJSUdI7bJ8PNJAg4lnl_Gg      4   

                                                                                                                                                                                                                                                                                                                                                                                                                                        
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- accept one document at a time
- return a list of tokens

You are free to use any method you have learned this week.

```python
# import spacy.cli
# spacy.cli.download("en_core_web_sm")
```

```python
# Optional: Consider using spaCy in your function. The spaCy library can be imported by running this cell.
# A pre-trained model (en_core_web_sm) has been made available to you in the CodeGrade container.
# If you DON'T need use the en_core_web_sm model, you can comment it out below.

import spacy
nlp = spacy.load('en_core_web_sm')
```

```python
import re
```

<!-- grade_id: cell-3d570d5a1cd6cb64, points: 0, grade: false, solution: true -->
```python
def tokenize(doc):
    doc = nlp(doc)
    return [token.lemma_.strip() for token in doc if (token.is_stop != True) and (token.is_punct != True)]
```

```python
df['tokens'] = df['text'].apply(tokenize)
df['tokens'].head(3)
```

Output:
```
0                                                                                                                                                      [beware, fake, fake, fake, small, business, Los, Alamitos, receive, look, like, legitimate, bill, $, 70, account, number, , call, phone, number, list, 866, 273, 7934, , wait, time, hold, say, 20, minute, leave, message, , live, person, phone, matter, number, select, , leave, firm, message, contact, BBB, attorney, company, try, scam, business, illegal]
1                                                                                                                                                                                                                                                                                                                                                                                                                       [come, lunch, Togo, Service, quick, staff, friendly, complaint, sweet, tea, good, parki
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
2. Write a fake review and query for the 10 most similiar reviews, print the text of the reviews. Do you notice any patterns?
    - Given the size of the dataset, use `NearestNeighbors` model for this. 

```python
""" Import Statements """

# Classics
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt

from sklearn.feature_extraction.text import CountVectorizer, TfidfVectorizer
from sklearn.metrics.pairwise import cosine_similarity
from sklearn.neighbors import NearestNeighbors
from sklearn.decomposition import PCA
```

```python
texts = df['text'].tolist()
```

```python
stars = df['stars'].tolist()
```

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
# Instantiate vectorizer object
tfidf = TfidfVectorizer(stop_words='english',
                        tokenizer=tokenize,
                        ngram_range=(1,2),
                        min_df=3,
                        max_df=0.25)

# Create a vocabulary and get word counts per document
dtm = tfidf.fit_transform(df['text'])

# Get feature names to use as dataframe column headers
dtm = pd.DataFrame(dtm.todense(), columns=tfidf.get_feature_names())
print(dtm.shape)
dtm.head()
```

Output:
```
/opt/anaconda3/lib/python3.8/site-packages/sklearn/feature_extraction/text.py:382: UserWarning: Your stop_words may be inconsistent with your preprocessing. Tokenizing the stop words generated tokens ['nt'] not in stop_words.
  warnings.warn('Your stop_words may be inconsistent with '
(10000, 32814)
          $    +   -duck   -i   -it   -the   -they    1   1/2  ...    à  à ce  \
0  0.0  0.0  0.0     0.0  0.0   0.0    0.0     0.0  0.0   0.0  ...  0.0   0.0   
1  0.0  0.0  0.0     0.0  0.0   0.0    0.0     0.0  0.0   0.0  ...  0.0   0.0   
2  0.0  0.0  0.0     0.0  0.0   0.0    0.0     0.0  0.0   0.0  ...  0.0   0.0   
3  0.0  0.0  0.0     0.0  0.0   0.0    0.0     0.0  0.0   0.0  ...  0.0   0.0   
4  0.0  0.0  0.0     0.0  0.0   0.0    0.0     0.0  0.0   0.0  ...  0.0   0.0   

   à essayer  à la   ça  équipe  était  été  être    ‍  
0        0.0   0.0  0.0     0.0    0.0  0.0   0.0  0.0  
1        0.0   0.0  0.0     0.0    0.0  0.0   0.0  0.0  
2        0.0   0.0  0.0     0.0    0.0  0
```

```python
og_df = df
og_df
```

Output:
```
business_id  cool                date  funny  \
0     nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1     eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2     6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3     k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4     6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   
...                      ...   ...                 ...    ...   
9995  1h3ysSuSazvXc1aeLiiOew     0 2017-10-07 10:57:15      1   
9996  Rwahe1zbFpw6VIjb5ngZeg     0 2014-01-18 15:52:52      0   
9997  8itGZAOBMiTbHKOwLuh4_Q     0 2018-08-26 02:53:21      0   
9998  A5Rkh7UymKm0_Rxm9K2PJw     0 2018-04-23 23:36:07      0   
9999  xtYiHTmunjfCN2sUaQxBjA     0 2016-06-02 23:37:51      0   

                   review_id  stars  \
0     eZs2tpEJtXPwawvHnHZIgQ      1   
1     DoQDWJsNbU0KL1O29l_Xug      4   
2     DDOdGU7zh56yQHmUnL1idQ      3   
3     LfTMUWnfGFMOfOIyJcwLVA      1   
4     zJSUdI7bJ8PNJAg4lnl_Gg      4   
..
```

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors
from sklearn.metrics.pairwise import cosine_similarity

# Calculate Distance of TF-IDF Vectors
dist_matrix = cosine_similarity(dtm)
print(dist_matrix.shape)

# Turn it into a DataFrame
matrix = pd.DataFrame(dist_matrix)
matrix.head()
```

Output:
```
(10000, 10000)
       0         1         2         3         4         5         6     \
0  1.000000  0.000000  0.019836  0.000000  0.005750  0.012194  0.014584   
1  0.000000  1.000000  0.021128  0.000000  0.020772  0.014510  0.000000   
2  0.019836  0.021128  1.000000  0.011037  0.011748  0.005236  0.014809   
3  0.000000  0.000000  0.011037  1.000000  0.008656  0.029411  0.009737   
4  0.005750  0.020772  0.011748  0.008656  1.000000  0.018331  0.000000   

       7         8         9     ...      9990      9991      9992      9993  \
0  0.000000  0.008414  0.015143  ...  0.000000  0.000000  0.009606  0.016769   
1  0.000000  0.000000  0.000000  ...  0.008868  0.010311  0.000000  0.000000   
2  0.000000  0.040473  0.019201  ...  0.015430  0.024717  0.009444  0.023924   
3  0.000000  0.007261  0.000000  ...  0.000000  0.000000  0.009903  0.000000   
4  0.016289  0.008700  0.043269  ...  0.006784  0.005050  0.031305  0.000000   

       9994      9995      9996      9997      9998  
```

```python
from sklearn.neighbors import NearestNeighbors

# Fit on DTM
nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree')
nn.fit(dtm)

# sample a doc from dtm to use as our query point 
doc = dtm.iloc[0].values

# Query Using kneighbors 
nn.kneighbors([doc])
```

Output:
```
(array([[0.        , 1.        , 1.        , 1.24099148, 1.28395851,
         1.2847304 , 1.30464156, 1.31772762, 1.31970901, 1.32708829]]),
 array([[   0, 6311, 6204, 2131, 6132, 3276, 7063, 9036, 2943, 6019]]))
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

fake_review = ['''Beautiful new location for Maman near NYU! The outdoor floral decorations are stunning. We sat next to the French windows that I believe open to a porch balcony vibe in the warmer months.

The staff was very good about keeping everyone socially distanced and taking peoples temperatures and contact information if contact tracing was needed.

My friend and I shared some cookies and baked pastries with the famous Lavendar hot chocolate and earl grey tea.

Great service will come again!''']

fake_review_vect = tfidf.transform(fake_review)
fake_review_vect
```

Output:
```
<1x32814 sparse matrix of type '<class 'numpy.float64'>'
	with 49 stored elements in Compressed Sparse Row format>
```

```python
review_nn = nn.kneighbors(fake_review_vect.todense())
review_nn
```

Output:
```
(array([[1.        , 1.        , 1.30515403, 1.31540454, 1.31634325,
         1.31713415, 1.32501192, 1.3254119 , 1.3296621 , 1.33030418]]),
 array([[6311, 6204, 6132, 4977, 2779, 4282, 1595, 4720, 5241, 6279]]))
```

```python
type(review_nn)
```

Output:
```
tuple
```

```python
review_nn[1][0]
```

Output:
```
array([6311, 6204, 6132, 4977, 2779, 4282, 1595, 4720, 5241, 6279])
```

```python
for review_num in review_nn[1][0]:
    print(review_num, texts[review_num])
    print()
```

Output:
```
6311 天氣很熱吃不下東西，今天我點了一個韓國冷面湯、餐後點了甜點，冰沙系列不會太甜膩，覺得店家很用心製作，包含擺盤精緻、佐料衛生，夏日想開胃，這是一個不錯的選擇，服務人員也很敬業，以後會常常來

6204 旅行でラスベガスに来ましたがネイルがはげてるのが気になり、探したお店でした。
質問にも丁寧に答えてくれましたし、日本人の方も日本語が話せる方も居て、とても綺麗で居心地のいいお店でした。 
ネイルはちはるさんと言う綺麗な方が丁寧にしてくれとても気に入りました。
予定になかったまつ毛エクステもお願いし、日本ではまだあまりないブラウンカラーのエクステをしてもらい、とても気に入りました。
また是非マッサージなどで伺いたいと思います。

6132 Hi, 
Update: I went to the location and spoke to the people around, they think there is something fishy going on...if someone is able to get some information, please contact me...my email id is akshayinfy@gmail.com

i am not able to contact these people since the last  few days, they have had my car for a while and are not replying to phone calls or emails...i see others also having this problem, does anyone have any contact numbers of any staff member? Someone named Brian owns the place...does anybody have any phone number for this person?????

4977 Very cute and cozy little coffee shop/bakery. They even have gelato. Prices are very reasonable and the two gentlemen
```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset. 

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels). 
    - Use that Pipeline to predict a star rating for your fake review from Part 2. 



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`. 
    - Include 2 possible values for each parameter
    - **Use `n_jobs` = 1** 
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import GridSearchCV

# Name the gridsearch instance "gs"
```

```python
from sklearn.svm import LinearSVC

# Create Pipeline Components
vect = TfidfVectorizer(stop_words='english', 
                       ngram_range=(1,2),# result lies in 1 word or 2 words
                       min_df=2,
                       max_df=0.25)
svm = LinearSVC()
```

```python
# Define the Pipeline
pipe = Pipeline([
    ('vect', vect),      # TF-IDF Vectorizer
    ('clf', svm)         # clf: classifier, LinearSVC Classifier
])
pipe
```

Output:
```
Pipeline(memory=None,
         steps=[('vect',
                 TfidfVectorizer(analyzer='word', binary=False,
                                 decode_error='strict',
                                 dtype=<class 'numpy.float64'>,
                                 encoding='utf-8', input='content',
                                 lowercase=True, max_df=0.25, max_features=None,
                                 min_df=2, ngram_range=(1, 2), norm='l2',
                                 preprocessor=None, smooth_idf=True,
                                 stop_words='english', strip_accents=None,
                                 sublinear_tf=False,
                                 token_pattern='(?u)\\b\\w\\w+\\b',
                                 tokenizer=None, use_idf=True,
                                 vocabulary=None)),
                ('clf',
                 LinearSVC(C=1.0, class_weight=None, dual=True,
                           fit_intercept=True, intercept_scaling=1,
          
```

```python
parameters = {
    'vect__max_df': (0.3, 0.5),
    'clf__C': (0.5, 1.)
}

gs = GridSearchCV(pipe, parameters, cv=5, n_jobs=1, verbose=1)
```

```python
df.head()
```

Output:
```
business_id  cool                date  funny  \
0  nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1  eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2  6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3  k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4  6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   

                review_id  stars  \
0  eZs2tpEJtXPwawvHnHZIgQ      1   
1  DoQDWJsNbU0KL1O29l_Xug      4   
2  DDOdGU7zh56yQHmUnL1idQ      3   
3  LfTMUWnfGFMOfOIyJcwLVA      1   
4  zJSUdI7bJ8PNJAg4lnl_Gg      4   

                                                                                                                                                                                                                                                                                                                                                                                                                                        
```

```python
# stars
```

```python
gs.fit(texts, stars)
```

Output:
```
Fitting 5 folds for each of 4 candidates, totalling 20 fits
[Parallel(n_jobs=1)]: Using backend SequentialBackend with 1 concurrent workers.
[Parallel(n_jobs=1)]: Done  20 out of  20 | elapsed:   46.1s finished
GridSearchCV(cv=5, error_score=nan,
             estimator=Pipeline(memory=None,
                                steps=[('vect',
                                        TfidfVectorizer(analyzer='word',
                                                        binary=False,
                                                        decode_error='strict',
                                                        dtype=<class 'numpy.float64'>,
                                                        encoding='utf-8',
                                                        input='content',
                                                        lowercase=True,
                                                        max_df=0.25,
                                                        max_fea
```

```python
gs.best_score_
```

Output:
```
0.6188
```

```python
gs.best_params_
```

Output:
```
{'clf__C': 0.5, 'vect__max_df': 0.5}
```

```python
pred = gs.predict(fake_review)
```

```python
pred
```

Output:
```
array([5])
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["I wish dogs knew how to speak English."])[0]
assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations. Refer to yesterday's notebook to extract. 
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this: 

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.** 

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz). 

```python
from gensim import corpora
# Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA 
from gensim.models.ldamodel import LdaModel
import gensim
import re
```

```python
df.head()
```

Output:
```
business_id  cool                date  funny  \
0  nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1  eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2  6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3  k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4  6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   

                review_id  stars  \
0  eZs2tpEJtXPwawvHnHZIgQ      1   
1  DoQDWJsNbU0KL1O29l_Xug      4   
2  DDOdGU7zh56yQHmUnL1idQ      3   
3  LfTMUWnfGFMOfOIyJcwLVA      1   
4  zJSUdI7bJ8PNJAg4lnl_Gg      4   

                                                                                                                                                                                                                                                                                                                                                                                                                                        
```

### 1. Estimate a LDA topic model of the review tex

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# Remember to read the LDA docs for more information on the various class attirbutes and methods available to you
# in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

# don't change this value 
num_topics = 5
# use tokenize function you created earlier to create tokens

# create a id2word object (hint: use corpora.Dictionary)
id2word = corpora.Dictionary(df['tokens'])
# create a corpus object (hint: id2word.doc2bow)
corpus = [id2word.doc2bow(text) for text in df['tokens']]
# instantiate an lda model
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )
```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 1-2 visualizations of the results

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
import seaborn as sns
import matplotlib.pyplot as plt

import pyLDAvis
import pyLDAvis.gensim 

# Use pyLDAvis (or a ploting tool of your choice) to visualize your results 

pyLDAvis.enable_notebook()
vis = pyLDAvis.gensim.prepare(lda, corpus, id2word)
vis
```

Output:
```
PreparedData(topic_coordinates=              x         y  topics  cluster       Freq
topic                                                
1     -0.122055  0.047167       1        1  32.398080
4      0.026664 -0.074342       2        1  29.136449
0     -0.065530 -0.018309       3        1  29.131240
3      0.056445 -0.037212       4        1   5.938803
2      0.104477  0.082697       5        1   3.395428, topic_info=        Term          Freq         Total Category  logprob  loglift
0             27343.000000  27343.000000  Default  30.0000  30.0000
39      time   4283.000000   4283.000000  Default  29.0000  29.0000
47      good   6516.000000   6516.000000  Default  28.0000  28.0000
130    place   5356.000000   5356.000000  Default  27.0000  27.0000
82        go   2968.000000   2968.000000  Default  26.0000  26.0000
..       ...           ...           ...      ...      ...      ...
354     find     58.552071   1610.528152   Topic5  -5.7109   0.0683
429      day     56.666271   1541.1
```

#### 3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

- topic 4 and topic 2 is a little overlapped, for example, word "warranty" is one overlapping word
- "good" is pretty popular in topic 1 and popular overall
- "milkshake" is most likely appeared in topic 3
- topic 1 and 3 are mostly likely food review while topic 2 in non-food review


## Stretch Goals

Complete one of more of these to push your score towards a three: 
* Incorporate named entity recognition into your analysis
* Compare vectorization methods in the classification section
* Analyze more (or all) of the yelp dataset - this one is v. hard. 
* Use a generator object on the reviews file - this would help you with the analyzing the whole dataset.
* Incorporate any of the other yelp dataset entities in your analysis (business, users, etc.)